```python
# Initialize Otter
import otter
grader = otter.Notebook("worksheet10_linear_regression_pytorch.ipynb")
```

![](img/logo.png)

# Worksheet 10: Build a Simple Linear Regression Model with PyTorch

## Exercise 1: Linear Regression with a Neural Network

{rubric: accuracy = 10}

Build and train a simple PyTorch model to learn GPA from hours of study.

A setup cell is provided below to generate the dataset and tensors. Complete the solution cell only.

### Requirements
1. Define a linear regression model named `model` (for example, one `nn.Linear` layer).
2. Use a regression loss function named `criterion` and an optimizer named `optimizer`.
3. Train the model for multiple epochs.
4. Store training predictions in `predictions` and print learned parameters and predicted GPA values.

```python
# Setup code (provided)
import torch

rng_seed = 90
torch.manual_seed(rng_seed)

# Dataset
X = [1.0, 2.0, 3.0, 4.0]
Y = [2.5, 3.0, 3.5, 4.0]

# Convert to tensors with shape (n, 1)
X_tensor = torch.tensor(X, dtype=torch.float32).view(-1, 1)
Y_tensor = torch.tensor(Y, dtype=torch.float32).view(-1, 1)

```

```python
...

```

Run the test below to see if you have done it correctly
> Note: There might be hidden tests that are not shown here

```python
grader.check("ex1")
```

## Exercise 2: Learning Rate and Convergence

{rubric: accuracy = 10}

In this exercise, compare how different learning rates affect convergence.

Use the setup code below. Then train one model per learning rate in `learning_rates` for `epochs` epochs and:

1. Store all loss curves in a dictionary `loss_histories` where keys are learning rates and values are lists of loss values per epoch.
2. Store final losses in a dictionary `final_losses`.
3. Set `best_lr` to the learning rate with the lowest final loss.

```python
# Setup code (provided)
import torch
import torch.nn as nn

torch.manual_seed(rng_seed)

X2_tensor = torch.tensor([1.0, 2.0, 3.0, 4.0], dtype=torch.float32).view(-1, 1)
Y2_tensor = torch.tensor([2.5, 3.0, 3.5, 4.0], dtype=torch.float32).view(-1, 1)

learning_rates = [0.001, 0.01, 0.1]
epochs = 300

```

```python
...

```

Run the test below to see if you have done it correctly
> Note: There might be hidden tests that are not shown here

```python
grader.check("ex2")
```

## Submission instructions

{rubric: mechanics = 5}

- Make sure the notebook can run from top to bottom without any error. Restart the kernel and run all cells.
- Submit to Gradescope

Tests for ex1 (points: 10)
```
>>> import torch
>>> import torch.nn as nn
>>> assert isinstance(X_tensor, torch.Tensor), 'X_tensor must be a torch.Tensor'
>>> assert isinstance(Y_tensor, torch.Tensor), 'Y_tensor must be a torch.Tensor'
>>> assert X_tensor.shape == (4, 1), f'Unexpected X_tensor shape: {X_tensor.shape}'
>>> assert Y_tensor.shape == (4, 1), f'Unexpected Y_tensor shape: {Y_tensor.shape}'
>>> assert isinstance(model, nn.Module), 'model must be a torch.nn.Module'
>>> assert isinstance(criterion, nn.MSELoss), 'criterion should be nn.MSELoss'
>>> assert isinstance(optimizer, torch.optim.Optimizer), 'optimizer must be a torch optimizer'
>>> assert isinstance(predictions, torch.Tensor), 'predictions must be a torch.Tensor'
>>> assert predictions.shape == (4, 1), f'Unexpected predictions shape: {predictions.shape}'
```


Tests for ex2 (points: 10)
```
>>> assert isinstance(loss_histories, dict), 'loss_histories must be a dictionary'
>>> assert isinstance(final_losses, dict), 'final_losses must be a dictionary'
>>> assert set(loss_histories.keys()) == set(learning_rates), 'loss_histories keys must match learning_rates'
>>> assert set(final_losses.keys()) == set(learning_rates), 'final_losses keys must match learning_rates'
>>> assert all((len(v) == epochs for v in loss_histories.values())), 'Each loss history must have length equal to epochs'
>>> assert best_lr in learning_rates, 'best_lr must be one of learning_rates'
```
