```python
# Initialize Otter
import otter
grader = otter.Notebook("template.ipynb")
```

![](img/logo.png)

# Assignment template

## Exercise 1: [Insert title here]

{rubric: accuracy = 5}

[Insert description here, e.g. write a function to sum two numbers]

```python
def sum(a, b):
    ...

```

Run the test below to see if you have done it correctly
> Note: There might be hidden tests that are not shown here

```python
grader.check("ex1")
```

## Exercise 2: [Insert title here]

{rubric: accuracy = 5}

[Insert description here, e.g. import iris.csv as a pandas dataframe called `df`]

```python
import pandas as pd
df = ...
```

Run the test below to see if you have done it correctly
> Note: There might be hidden tests that are not shown here

```python
grader.check("ex2")
```

Run this code cell below to plot your data

```python
import matplotlib.pyplot as plt
# plot scatter plot
plt.scatter(df["sepal.length"], df["sepal.width"])
```

## Exercise 3 (MANUAL QUESTION): [Insert title here]

{rubric: reasoning = 5}

[Insert description here, e.g. import iris.csv as a pandas dataframe called `df`]

<!-- BEGIN QUESTION -->

**Question:** What do you notice about the scatter plot?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

## Submission instructions

{rubric: mechanics = 5}

- Make sure the notebook can run from top to bottom without any error. Restart the kernel and run all cells.
- Submit to Gradescope

Tests for ex1 (points: 5)
```
>>> assert sum(1, 2) == 3, f'unexpected {sum(1, 2)}'
```


Tests for ex2 (points: 5)
```
>>> assert df is not None, 'df does not exist'
```
