```python
# Initialize Otter
import otter
grader = otter.Notebook("worksheet10_linear_regression_pytorch.ipynb")
```

![](img/logo.png)

# Worksheet 10: Build a Simple Linear Regression Model with PyTorch

## Exercise 1: Linear Regression with a Neural Network

{rubric: accuracy = 10}

Build and train a simple PyTorch model to learn GPA from hours of study.

A setup cell is provided below to generate the dataset and tensors. Complete the solution cell only.

### Requirements
1. Define a linear regression model named `model` (for example, one `nn.Linear` layer).
2. Use a regression loss function named `criterion` and an optimizer named `optimizer`.
3. Train the model for multiple epochs.
4. Store training predictions in `predictions` and print learned parameters and predicted GPA values.

```python
# Setup code (provided)
import torch

rng_seed = 42
torch.manual_seed(rng_seed)

# Dataset
X = [1.0, 2.0, 3.0, 4.0]
Y = [2.5, 3.0, 3.5, 4.0]

# Convert to tensors with shape (n, 1)
X_tensor = torch.tensor(X, dtype=torch.float32).view(-1, 1)
Y_tensor = torch.tensor(Y, dtype=torch.float32).view(-1, 1)

```

```python
# BEGIN SOLUTION
import torch.nn as nn

# Define model, loss, and optimizer
model = nn.Linear(1, 1)
criterion = nn.MSELoss()
optimizer = torch.optim.SGD(model.parameters(), lr=0.01)

# Train
epochs = 2000
for _ in range(epochs):
    optimizer.zero_grad()
    y_pred = model(X_tensor)
    loss = criterion(y_pred, Y_tensor)
    loss.backward()
    optimizer.step()

# Final predictions
with torch.no_grad():
    predictions = model(X_tensor)

weight = model.weight.item()
bias = model.bias.item()
print(f"Learned weight: {weight:.4f}, bias: {bias:.4f}")
print("Predicted GPA values:", predictions.view(-1).tolist())
# END SOLUTION

```

Run the test below to see if you have done it correctly
> Note: There might be hidden tests that are not shown here

```python
grader.check("ex1")
```

## Exercise 2: Learning Rate and Convergence

{rubric: accuracy = 10}

In this exercise, compare how different learning rates affect convergence.

Use the setup code below. Then train one model per learning rate in `learning_rates` for `epochs` epochs and:

1. Store all loss curves in a dictionary `loss_histories` where keys are learning rates and values are lists of loss values per epoch.
2. Store final losses in a dictionary `final_losses`.
3. Set `best_lr` to the learning rate with the lowest final loss.

```python
# Setup code (provided)
import torch
import torch.nn as nn

torch.manual_seed(rng_seed)

X2_tensor = torch.tensor([1.0, 2.0, 3.0, 4.0], dtype=torch.float32).view(-1, 1)
Y2_tensor = torch.tensor([2.5, 3.0, 3.5, 4.0], dtype=torch.float32).view(-1, 1)

learning_rates = [0.001, 0.01, 0.1]
epochs = 300

```

```python
# BEGIN SOLUTION
loss_histories = {}
final_losses = {}

for lr in learning_rates:
    torch.manual_seed(rng_seed)
    model_lr = nn.Linear(1, 1)
    criterion_lr = nn.MSELoss()
    optimizer_lr = torch.optim.SGD(model_lr.parameters(), lr=lr)

    losses = []
    for _ in range(epochs):
        optimizer_lr.zero_grad()
        pred_lr = model_lr(X2_tensor)
        loss_lr = criterion_lr(pred_lr, Y2_tensor)
        loss_lr.backward()
        optimizer_lr.step()
        losses.append(loss_lr.item())

    loss_histories[lr] = losses
    final_losses[lr] = losses[-1]

best_lr = min(final_losses, key=final_losses.get)
print("Final losses by learning rate:", final_losses)
print("Best learning rate:", best_lr)
# END SOLUTION

```

Run the test below to see if you have done it correctly
> Note: There might be hidden tests that are not shown here

```python
grader.check("ex2")
```

## Submission instructions

{rubric: mechanics = 5}

- Make sure the notebook can run from top to bottom without any error. Restart the kernel and run all cells.
- Submit to Gradescope

Tests for ex1 (points: 10)
```
>>> import torch
>>> import torch.nn as nn
>>> assert isinstance(X_tensor, torch.Tensor), 'X_tensor must be a torch.Tensor'
>>> assert isinstance(Y_tensor, torch.Tensor), 'Y_tensor must be a torch.Tensor'
>>> assert X_tensor.shape == (4, 1), f'Unexpected X_tensor shape: {X_tensor.shape}'
>>> assert Y_tensor.shape == (4, 1), f'Unexpected Y_tensor shape: {Y_tensor.shape}'
>>> assert isinstance(model, nn.Module), 'model must be a torch.nn.Module'
>>> assert isinstance(criterion, nn.MSELoss), 'criterion should be nn.MSELoss'
>>> assert isinstance(optimizer, torch.optim.Optimizer), 'optimizer must be a torch optimizer'
>>> assert isinstance(predictions, torch.Tensor), 'predictions must be a torch.Tensor'
>>> assert predictions.shape == (4, 1), f'Unexpected predictions shape: {predictions.shape}'
```
```
>>> with torch.no_grad():
...     train_mse = torch.mean((model(X_tensor) - Y_tensor) ** 2).item()
>>> assert train_mse < 0.03, f'Model fit is too weak. MSE={train_mse:.5f}'
>>> with torch.no_grad():
...     pred_5 = model(torch.tensor([[5.0]], dtype=torch.float32)).item()
>>> assert 4.2 <= pred_5 <= 4.8, f'Prediction at x=5.0 is out of expected range: {pred_5:.4f}'
```


Tests for ex2 (points: 10)
```
>>> assert isinstance(loss_histories, dict), 'loss_histories must be a dictionary'
>>> assert isinstance(final_losses, dict), 'final_losses must be a dictionary'
>>> assert set(loss_histories.keys()) == set(learning_rates), 'loss_histories keys must match learning_rates'
>>> assert set(final_losses.keys()) == set(learning_rates), 'final_losses keys must match learning_rates'
>>> assert all((len(v) == epochs for v in loss_histories.values())), 'Each loss history must have length equal to epochs'
>>> assert best_lr in learning_rates, 'best_lr must be one of learning_rates'
```
```
>>> for lr in learning_rates:
...     assert loss_histories[lr][-1] < loss_histories[lr][0], f'Loss did not decrease for lr={lr}'
>>> expected_best = min(final_losses, key=final_losses.get)
>>> assert best_lr == expected_best, 'best_lr should correspond to the minimum final loss'
>>> assert final_losses[0.1] < final_losses[0.001], 'Higher learning rate should converge faster here within fixed epochs'
>>> assert final_losses[best_lr] < 0.05, f'Best final loss is too high: {final_losses[best_lr]:.5f}'
```
